<problem display_name="Exercise 5" group_access="{&quot;51&quot;: [1, 2]}" markdown="null">
  <text>
    <p>
      <b>ESTIMATED TIME TO COMPLETE: 4 minutes</b>
    </p>
    <script type="text/python" system_path="python_lib">
def test_list(expect, ans):
  # Check that expect &amp; ans are lists
  assert (expect[0] == '[' and expect[-1] == ']')
  if '[' not in ans or ']' not in ans: return False

  expect = list( expect.replace(' ','')[1:-1].split(',') )
  ans = list( ans.replace(' ','')[1:-1].split(',') )
  return len(expect) == len(ans) and not (False in [e == a for (e,a) in zip(expect, ans)])
</script>
    <p>
Here is a different piece of code for working with lists: </p>
    <pre>
def applyEachTo(L, x):
    result = []
    for i in range(len(L)):
        result.append(L[i](x))
    return result
</pre>
    <p>
Suppose that you are given the following functions: </p>
    <pre>
def square(a):
    return a*a

def halve(a):
    return a/2

def inc(a):
    return a+1
</pre>
    <p>
For each of the following questions, indicate what value is returned. If you believe that an error will occur, write the word 'error'. </p>
    <ol class="enumerate">
      <li>
        <pre>
applyEachTo([inc, square, halve, abs], -3)
 </pre>
        <p>
          <customresponse cfn="test_list" expect="[-2, 9, -1.5, 3]">
            <textline size="50" correct_answer="[-2, 9, -1.5, 3]"/>
          </customresponse>
        </p>
      </li>
      <li>
        <pre>
applyEachTo([inc, square, halve, abs], 3.0)
 </pre>
        <p>
          <customresponse cfn="test_list" expect="[4.0, 9.0, 1.5, 3.0]">
            <textline size="50" correct_answer="[4.0, 9.0, 1.5, 3.0]"/>
          </customresponse>
        </p>
      </li>
      <li>
        <pre>
applyEachTo([inc, max, int], -3)
 </pre>
        <p>
          <stringresponse answer="error">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
    </ol>
  </text>
</problem>
