<problem display_name="Exercise 3" group_access="{&quot;51&quot;: [1, 2]}" markdown="null">
  <text>
    <p>
      <b>ESTIMATED TIME TO COMPLETE: 13 minutes</b>
      <br/>
      <b>      Note that you will have to answer all questions before you can click the Check button.</b>
    </p>
    <script type="text/python" system_path="python_lib">
def test_str(expect, ans):
  # Strip away ' ' or " " wrapping string input answer
  # (check which, if any, are present - don't want to pull valid
  #  characters out of the center of a string, eg "That's crazy!",
  #  nor do we wish to accept an input with mismatched modifiers, 
  #  eg 'abc")
  if len(expect) == 0 or len(ans) == 0: return expect == ans

  if ans[0] == "'" and ans[-1] == "'": ans = ans.strip("'")
  elif ans[0] == '"' and ans[-1] == '"': ans = ans.strip('"')

  if expect[0] == "'": expect = expect.strip("'")
  elif expect[0] == '"': expect = expect.strip('"')

  return expect == ans

def test_list_str(expect, ans):
  # For a list that can be comprised of strings
  # Check that expect &amp; ans are lists
  assert (expect[0] == '[' and expect[-1] == ']')
  if '[' not in ans or ']' not in ans: return False

  expect = list( expect.replace(' ','')[1:-1].split(',') )
  ans = list( ans.replace(' ','')[1:-1].split(',') )
  return len(expect) == len(ans) and not (False in [test_str(e,a) for (e,a) in zip(expect, ans)])
</script>
    <p>For each of the expressions below, specify its type and value. If it generates an error, 
      select type 'NoneType' and put the word 'error' in the box for the value. If it would be a function, 
      select type 'function' and put the word 'function' in the box for the value. </p>
    <p>If the method returns <code>None</code>, select type 'NoneType' and put the word 'None' in the box for the value. </p>
    <p>Assume we've made the following assignments: </p>
    <pre>
&gt; listA = [1, 4, 3, 0]
&gt; listB = ['x', 'z', 't', 'q']
</pre>
    <p> You may want to refer to the <a href="http://docs.python.org/tutorial/datastructures.html" target="_blank">Python Library Reference</a> 
      to learn about list methods. Assume these calls appear one after another (so if the list is modified in a question, 
      that modification stays for subsequent questions). </p>
    <ol class="enumerate">
      <li>
        <p>
          <code>listA.sort</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="function"/>
          </optionresponse>
          <stringresponse answer="function">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.sort()</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="None">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="list"/>
          </optionresponse>
          <customresponse cfn="test_list_str" expect="[0, 1, 3, 4]">
            <textline size="50" correct_answer="[0, 1, 3, 4]"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.insert(0, 100)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="None">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.remove(3)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="None">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.append(7)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="None">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="list"/>
          </optionresponse>
          <customresponse cfn="test_list_str" expect="[100, 0, 1, 4, 7]">
            <textline size="50" correct_answer="[100, 0, 1, 4, 7]"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA + listB</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="list"/>
          </optionresponse>
          <customresponse cfn="test_list_str" expect="[100, 0, 1, 4, 7, 'x', 'z', 't', 'q']">
            <textline size="50" correct_answer="[100, 0, 1, 4, 7, 'x', 'z', 't', 'q']"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listB.sort()</code>
          <br/>
          <code>listB.pop()</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'z'">
            <textline size="50" correct_answer="'z'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listB.count('a')</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="int"/>
          </optionresponse>
          <stringresponse answer="0">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listB.remove('a')</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="error">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.extend([4, 1, 6, 3, 4])</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="None">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.count(4)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="int"/>
          </optionresponse>
          <stringresponse answer="3">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.index(1)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="int"/>
          </optionresponse>
          <stringresponse answer="2">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.pop(4)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="int"/>
          </optionresponse>
          <stringresponse answer="7">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA.reverse()</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="None">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>listA</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','function','string','tuple','list')" correct="list"/>
          </optionresponse>
          <customresponse cfn="test_list_str" expect="[4, 3, 6, 1, 4, 4, 1, 0, 100]">
            <textline size="50" correct_answer="[4, 3, 6, 1, 4, 4, 1, 0, 100]"/>
          </customresponse>
        </p>
      </li>
    </ol>
  </text>
</problem>
