<problem display_name="Exercise 2 part 2" group_access="{&quot;51&quot;: [1, 2]}" markdown="null" max_attempts="9999" weight="0.0">
  <text>
    <p>
      <b>ESTIMATED TIME TO COMPLETE: 7 minutes</b>
      <br/>
      <b>      Note that you will have to answer all questions before you can click the Check button.</b>
    </p>
    <script type="text/python" system_path="python_lib">
def test_str(expect, ans):
  # Strip away ' ' or " " wrapping string input answer
  # (check which, if any, are present - don't want to pull valid
  #  characters out of the center of a string, eg "That's crazy!",
  #  nor do we wish to accept an input with mismatched modifiers, 
  #  eg 'abc")

  if ans[0] == "'" and ans[-1] == "'": ans = ans.strip("'")
  elif ans[0] == '"' and ans[-1] == '"': ans = ans.strip('"')

  if expect[0] == "'": expect = expect.strip("'")
  elif expect[0] == '"': expect = expect.strip('"')

  return expect == ans
</script>
    <p>
For each of the expressions below, specify its type and value. If it generates an error, select type 'NoneType' and write the word 'error' (note this is a word, not a string, no quotes) in the box for the value. While you could simply type these expressions into your IDE, we encourage you to answer them directly since this will help reinforce your understand of basic Python expressions. </p>
    <p>
Assume we've made the following assignments: </p>
    <pre>
&gt; str1 = 'hello'
&gt; str2 = ','
&gt; str3 = 'world'
</pre>
    
      <p>
      <section class="hints">
        <div class="collapsible">
          <header>
            <a href="#" id="id72">Note: Advanced String Slicing</a>
          </header>
          <section id="id71">
            <p>You've seen in lecture that you can slice a string with a call such as <code>s[i:j]</code>, which gives you a portion of string <code>s</code> from index <code>i</code>
	    to index <code>j-1</code>. However this is not the only way to slice a string! If you omit the starting index, Python will assume that you wish to start your slice
	      at index 0. If you omit the ending index, Python will assume you wish to end your slice at the end of the string. Check out this session with the Python shell:</p>
            <pre>
&gt;&gt;&gt; s = 'Python is Fun!'
&gt;&gt;&gt; s[1:5]
'ytho'
&gt;&gt;&gt; s[:5]
'Pytho'
&gt;&gt;&gt; s[1:]
'ython is Fun!'
&gt;&gt;&gt; s[:]
'Python is Fun!'
</pre>
            <p>That last example is interesting! If you omit both the start and ending index, you get your original string!</p>
            <p>There's one other cool thing you can do with string slicing. You can add a third parameter, <code>k</code>, like this: <code>s[i:j:k]</code>. This gives a slice
	    of the string <code>s</code> from index <code>i</code> to index <code>j-1</code>, with step size <code>k</code>. Check out the following examples:</p>
            <pre>
&gt;&gt;&gt; s = 'Python is Fun!'
&gt;&gt;&gt; s[1:12:2]
'yhni u'
&gt;&gt;&gt; s[1:12:3]
'yoiF'
&gt;&gt;&gt; s[::2]
'Pto sFn'
</pre>
            <p>The last example is similar to the example <code>s[:]</code>. With <code>s[::2]</code>, we're asking for the full string <code>s</code> (from index 0 through 13), 
	    with a step size of 2 - so we end up with every other character in <code>s</code>. Pretty cool!</p>
          </section>
        </div>
      </section>
    </p>
    
    
    <p>
      <section class="hints">
        <div class="collapsible">
          <header>
            <a href="#" id="id71">Note: The Python 'in' operator</a>
          </header>
          <section id="id71">
            <p>The operators <code>in</code> and <code>not in</code> test for collection membership (a 'collection' refers to a string, list, tuple or dictionary - don't worry, we will cover lists, tuples and dictionaries soon!). The expression </p>
            <pre>
element in coll
</pre>
            <p>
 evaluates to <code>True</code> if <code>element</code> is a member of the collection <code>coll</code>, and <code>False</code> otherwise. </p>
            <p>
The expression </p>
            <pre>
element not in coll
</pre>
            <p>
 evaluates to <code>True</code> if <code>element</code> is <b class="bf">not</b> a member of the collection <code>coll</code>, and <code>False</code> otherwise. </p>
            <p>
Note this returns the negation of <code>element in coll</code> - that is, the expression <code>element not in coll</code>        is equivalent to the expression <code>not (element in coll)</code>. </p>
          </section>
        </div>
      </section>
    </p>
    <ol class="enumerate">
      <li>
        <p>
          <code>str1[len(str1)]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="error">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1 + str2 + str3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'hello,world'">
            <textline size="50" correct_answer="'hello,world'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1 + str2 + ' ' + str3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'hello, world'">
            <textline size="50" correct_answer="'hello, world'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str3 * 3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'worldworldworld'">
            <textline size="50" correct_answer="'worldworldworld'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>'hello' == str1</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="boolean"/>
          </optionresponse>
          <stringresponse answer="True">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>

    </ol>
  </text>
</problem>
